// Numbas version: finer_feedback_settings
{"name": "potenser_allespm", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "potenser_allespm", "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "", "advice": "<p><strong>Niv&aring; 2.2</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Husk:</p>\n<ul>\n<li>++ blir +</li>\n<li>--&nbsp; &nbsp;blir +</li>\n<li>+-&nbsp; blir -</li>\n<li>-+&nbsp; blir -</li>\n</ul>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p><strong>Niv&aring; 3.1</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p>Husk:</p>\n<ul>\n<li>++ blir +</li>\n<li>--&nbsp; &nbsp;blir +</li>\n<li>+-&nbsp; blir -</li>\n<li>-+&nbsp; blir -</li>\n</ul>\n<p>N&aring;r en potens er opph&oslash;yd igjen med en annen eksponent s&aring; regnes dette ut ved &aring; multiplisere eksponentene:<br/>$\\left(a^b\\right)^c=a^{b\\cdot c}$</p>\n<p><strong>Niv&aring; 3.2</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p>N&aring;r en potens er opph&oslash;yd igjen med en annen eksponent s&aring; regnes dette ut ved &aring; multiplisere eksponentene:<br/>$\\left(a^b\\right)^c=a^{b\\cdot c}$</p>\n<p>Dersom vi har felles grunntall i potensene, s&aring; kan vi benytte oss av f&oslash;lgende regneregel for potens:<br/>$a^n\\cdot a^m=a^{n+m}$</p>\n<p>Husk:</p>\n<ul>\n<li>++ blir +</li>\n<li>--&nbsp; &nbsp;blir +</li>\n<li>+-&nbsp; blir -</li>\n<li>-+&nbsp; blir -</li>\n</ul>\n<p><strong>Niv&aring; 2.3</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p>N&aring;r man multipliserer ett tall med en potens, m&aring; man huske p&aring; regnerekkef&oslash;lgen</p>\n<ol>\n<li>Parenteser</li>\n<li>Potenser</li>\n<li>Multiplikasjon og Divisjon</li>\n<li>Addisjon og Subtraksjon</li>\n</ol>\n<p><strong>Niv&aring; 1.2</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p>N&aring;r man multipliserer ett tall med en potens, m&aring; man huske p&aring; regnerekkef&oslash;lgen</p>\n<ol>\n<li>Parenteser</li>\n<li>Potenser</li>\n<li>Multiplikasjon og Divisjon</li>\n<li>Addisjon og Subtraksjon</li>\n</ol>\n<p><strong>Niv&aring; 2.1</strong>&nbsp;</p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>N&aring;r vi har en br&oslash;k som er opph&oslash;yd med en eksponent, s&aring; kan vi ta denne faktoren p&aring; hver av komponentene inne i br&oslash;ken:<br/>$\\left(\\dfrac{a}{b}\\right)^c=\\dfrac{a^c}{b^c}$</p>\n<p>Husk:&nbsp;</p>\n<ul>\n<li>++ blir +</li>\n<li>--&nbsp; &nbsp;blir +</li>\n<li>+-&nbsp; blir -</li>\n<li>-+&nbsp; blir -</li>\n</ul>\n<p>Dette gjelder ogs&aring; dersom vi har negative tall i teller og/eller nevner!</p>\n<p><strong>Niv&aring; 1.1</strong></p>\n<p>Potens er definert som $a^3=a\\cdot a \\cdot a$ (alts&aring; $a$ multiplisert med seg selv $3$ ganger)</p>\n<p>Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</p>\n<p>Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</p>\n<p>Dersom vi har negativ eksponent, s&aring; kan vi uttrykke potensen som en br&oslash;k:<br/>$a^{-b}=\\dfrac{1}{a^b}$</p>\n<p>Husk:</p>\n<ul>\n<li>++ blir +</li>\n<li>--&nbsp; &nbsp;blir +</li>\n<li>+-&nbsp; blir -</li>\n<li>-+&nbsp; blir -</li>\n</ul>\n<p>N&aring;r en potens er opph&oslash;yd igjen med en annen eksponent s&aring; regnes dette ut ved &aring; multiplisere eksponentene:<br/>$\\left(a^b\\right)^c=a^{b\\cdot c}$</p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a22": {"name": "a22", "group": "Niv\u00e5 2.2", "definition": "random(-5..5 except 1 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b22": {"name": "b22", "group": "Niv\u00e5 2.2", "definition": "random(-3..-1)", "description": "", "templateType": "anything", "can_override": false}, "a31": {"name": "a31", "group": "Niv\u00e5 3.1", "definition": "random(-5..5)", "description": "", "templateType": "anything", "can_override": false}, "b31": {"name": "b31", "group": "Niv\u00e5 3.1", "definition": "random(-2..0)", "description": "", "templateType": "anything", "can_override": false}, "c31": {"name": "c31", "group": "Niv\u00e5 3.1", "definition": "random(-2..2)", "description": "", "templateType": "anything", "can_override": false}, "a32": {"name": "a32", "group": "Niv\u00e5 3.2", "definition": "random(-5..5 except 0 except 1 except -1)", "description": "", "templateType": "anything", "can_override": false}, "b32": {"name": "b32", "group": "Niv\u00e5 3.2", "definition": "random(1..4)", "description": "", "templateType": "anything", "can_override": false}, "c32": {"name": "c32", "group": "Niv\u00e5 3.2", "definition": "a32^2", "description": "", "templateType": "anything", "can_override": false}, "d32": {"name": "d32", "group": "Niv\u00e5 3.2", "definition": "random(-3..-1)", "description": "", "templateType": "anything", "can_override": false}, "a23": {"name": "a23", "group": "Niv\u00e523", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b23": {"name": "b23", "group": "Niv\u00e523", "definition": "random(-5..5 except 0 except -1 except 1)", "description": "", "templateType": "anything", "can_override": false}, "c23": {"name": "c23", "group": "Niv\u00e523", "definition": "random(-3..0)", "description": "", "templateType": "anything", "can_override": false}, "a12": {"name": "a12", "group": "Niv\u00e512", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b12": {"name": "b12", "group": "Niv\u00e512", "definition": "random(-5..5 except 0 except -1 except 1)", "description": "", "templateType": "anything", "can_override": false}, "c12": {"name": "c12", "group": "Niv\u00e512", "definition": "random(0..3)", "description": "", "templateType": "anything", "can_override": false}, "a21": {"name": "a21", "group": "Niv\u00e5 2.1", "definition": "random(-5..5)*b21", "description": "", "templateType": "anything", "can_override": false}, "b21": {"name": "b21", "group": "Niv\u00e5 2.1", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c21": {"name": "c21", "group": "Niv\u00e5 2.1", "definition": "random(0..3 except 1)", "description": "", "templateType": "anything", "can_override": false}, "a11": {"name": "a11", "group": "Niv\u00e5 1.1", "definition": "random(1..5)", "description": "", "templateType": "anything", "can_override": false}, "b11": {"name": "b11", "group": "Niv\u00e5 1.1", "definition": "random(1..5)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Niv\u00e5 2.2", "variables": ["a22", "b22"]}, {"name": "Niv\u00e5 3.1", "variables": ["a31", "b31", "c31"]}, {"name": "Niv\u00e5 3.2", "variables": ["a32", "b32", "c32", "d32"]}, {"name": "Niv\u00e523", "variables": ["a23", "b23", "c23"]}, {"name": "Niv\u00e512", "variables": ["a12", "b12", "c12"]}, {"name": "Niv\u00e5 2.1", "variables": ["a21", "b21", "c21"]}, {"name": "Niv\u00e5 1.1", "variables": ["a11", "b11"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 2", "marks": "56", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Sp\u00f8rsm\u00e5l 2", "rawLabel": "Sp\u00f8rsm\u00e5l 2", "otherPart": 1, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Sp\u00f8rsm\u00e5l 2", "rawLabel": "Sp\u00f8rsm\u00e5l 2", "otherPart": 4, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:</p>\n<p><span data-jme-visible=\"a22&lt;0\">$($</span>$\\var{a22}$<span data-jme-visible=\"a22&lt;0\">$)$</span>$^{\\var{b22}}$ (skriv svar som heltall eller br&oslash;k - eks. 2/3)</p>", "minValue": "a22^b22", "maxValue": "a22^b22", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 3.1", "marks": "29", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Sp\u00f8rsm\u00e5l 3", "rawLabel": "Sp\u00f8rsm\u00e5l 3", "otherPart": 2, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Sp\u00f8rsm\u00e5l 3", "rawLabel": "Sp\u00f8rsm\u00e5l 3", "otherPart": 3, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:&nbsp;</p>\n<p>$\\left(\\var{a31}^\\var{b31}\\right)^\\var{c31}$</p>\n<p></p>", "minValue": "(a31^b31)^c31", "maxValue": "(a31^b31)^c31", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 3.2", "marks": "15", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:&nbsp;</p>\n<p>$(\\var{a32})^{\\var{b32}} \\cdot (\\var{c32})^{\\var{d32}}$ (skriv svar som heltall eller br&oslash;k - eks. 2/3)</p>\n<p></p>", "minValue": "(a32^b32)*(c32^d32)", "maxValue": "(a32^b32)*(c32^d32)", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 2.3", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:</p>\n<p>$\\var{a23}\\cdot$<span data-jme-visible=\"b23&lt;0\">$($</span>$\\var{b23}$<span data-jme-visible=\"b23&lt;0\">$)$</span>$^\\var{c23}$ (skriv svar som heltall eller br&oslash;k - eks. 2/3)</p>", "minValue": "a23*(b23^c23)", "maxValue": "a23*(b23^c23)", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 1.2", "marks": "28", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Sp\u00f8rsm\u00e5l 3", "rawLabel": "Sp\u00f8rsm\u00e5l 3", "otherPart": 5, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Sp\u00f8rsm\u00e5l 3", "rawLabel": "Sp\u00f8rsm\u00e5l 3", "otherPart": 6, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:&nbsp;</p>\n<p>$\\var{a12}\\cdot$<span data-jme-visible=\"b12&lt;0\">$($</span>$\\var{b12}$<span data-jme-visible=\"b12&lt;0\">$)$</span>$^\\var{c12}$</p>", "minValue": "a12*(b12^c12)", "maxValue": "a12*(b12^c12)", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 2.1", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:&nbsp;</p>\n<p>$\\left( \\dfrac{\\var{a21}}{\\var{b21}}\\right)^\\var{c21}$&nbsp; (skriv svar som heltall eller br&oslash;k - eks. 2/3)</p>", "minValue": "(a21/b21)^c21", "maxValue": "(a21/b21)^c21", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Niv\u00e5 1.1", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Regn ut:&nbsp;</p>\n<p><span data-jme-visible=\"a11&lt;0\">$($</span>$\\var{a11}$<span data-jme-visible=\"a11&lt;0\">$)$</span>$^\\var{b11}$</p>", "minValue": "a11^b11", "maxValue": "a11^b11", "correctAnswerFraction": true, "allowFractions": true, "mustBeReduced": true, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "information", "useCustomName": true, "customName": "Slut 1", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Bra jobbat du &auml;r p&aring; en bra niv&aring; och nu &auml;r testen klar.&nbsp;</p>"}], "partsMode": "explore", "maxMarks": "100", "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false, "contributors": [{"name": "Lars Hansson", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/774/"}, {"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "resources": []}]}], "contributors": [{"name": "Lars Hansson", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/774/"}, {"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}]}